---
title       : Should Fail Build Test
description : Insert the chapter description here


--- type:NormalExercise lang:python xp:100 skills:2 key:43c6dd9bec
## I am a failure

asdf

*** =instructions
asdf

*** =hint
asdf

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}
x = 4
```

*** =solution
```{python}
x = 4
```

*** =sct
```{python}
test_object("y")
```
