---
title       : Should Fail Build Test
description : Insert the chapter description here


--- type:NormalExercise lang:python xp:100 skills:2 key:43c6dd9bec
## I am a failure

asdf

*** =instructions
asdf

*** =hint
asdf

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}
x = 4
```

*** =solution
```{python}
x = 4
```

*** =sct
```{python}
test_object("y")
```

--- type:NormalExercise lang:python xp:100 skills:2 key:8b1a3d29e9
## Converter Fail Build 


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}
from bokeh.plotting import figure
```

*** =sample_code
```{python}
p = figure()
```

*** =solution
```{python}
p = figure()
```

*** =sct
```{python}
test_object('p')
```
